---
title: Modeling
description: ""
---

## Baseline Model

```yaml
type: VideoExercise
key: 50c7b3ff4d
xp: 50
```

`@projector_key`
38086466ef7244576a066cefcf1affa2

---

## Example coding exercise

```yaml
type: NormalExercise
key: e8c1edbe67
lang: python
xp: 100
skills: 2
```

This is an example exercise.

`@instructions`


`@hint`


`@pre_exercise_code`
```{python}

```

`@sample_code`
```{python}

```

`@solution`
```{python}

```

`@sct`
```{python}

```
